```python
# Initialize Otter
import otter
grader = otter.Notebook("week1_introduction.ipynb")
```

<h2> Introduction to using Jupyter </h2>

This notebook is meant to give you an introduction to using Jupyter.

This notebook has a couple of questions that introduce you to the process of using Python code to answer questions, along with the particular syntax. If you're interested in learning more about the details of coding in Python, [this tutorial](https://www.w3schools.com/python/) is a great starting point.

Let's get started!

<h3> Part 1: Computing a sum with Python </h3>

In previous courses, you probability computed sums with sigma notation; for example,
$$\sum_{n = 1}^{1000} \frac{1}{n} = \frac 1 1 + \frac 1 2 + \frac 1 3 + \cdots \frac{1}{1000}.$$
These can get extremely cumbersome to compute by hand, but are easily handled by a computer using a [for loop](https://www.w3schools.com/python/python_for_loops.asp). The Python code that will compute this sum is below, with a lot of comments. Comments are placed on lines using \#, and are not interpreted as code.

```python
# initialize the sum
def sum_of_reciprocals(N):
    """Generate the sum of the first N recipricals, for n=1,...,N"""
    # initialize the sum
    s = 0
    # run a for loop that adds 1/n each
    for n in range(1, N + 1):
        s = s + 1/n
    return s
print(sum_of_reciprocals(1000))
```

### Problem 1a

So the sum of the first $1000$ numbers is about $7.49$. Now it's your turn; **answer the following two questions**.

* Create a function that calculates $\sum\limits_{n = 1}^{N} n^3$. Note that in Python, $a^b$ is written as `a**b`.



<!-- BEGIN QUESTION -->



```python
# Put your answer to Problem 1a here
def sum_of_cubes(N):
    """Generate the sum of the first N cubes, for n=1,...,N"""
    ...
    return s
```

```python
grader.check("q1a")
```

<!-- END QUESTION -->

### Problem 1b


Now create a function that calculates $\sum\limits_{n = 1}^{N} \frac{1}{n^2}$.

<!-- BEGIN QUESTION -->



```python
# Put your answer to Problem 1a here
def sum_of_square_inverses(N):
    """Generate the sum of the first N inverse of squares, for n=1,...,N"""
    ...
    return s
print(sum_of_square_inverses(20))
print(sum_of_square_inverses(100))
```

```python
grader.check("q1b")
```

<!-- END QUESTION -->

<h3> Part 2: Generating random numbers with Python </h3>

For this, we will use the [numpy](https://numpy.org) package, although there are many choices for generating random numbers using python.

Here is the documentation for the [randint()](https://numpy.org/doc/stable/reference/random/generated/numpy.random.randint.html) function we will be using.

```python
import numpy as np
# choose either 0 or 1
flip = np.random.randint(low = 0 , high = 2)
print(flip)
```

For this problem, we will think of 0 as Tails and 1 as Heads

```python
flips = np.random.randint(low = 0, high = 2, size = 100)
num_heads = flips.sum()
print(f'Number of heads = {num_heads}')
```

### Problem 2

Suppose you flip $1000$ coins and record the number of heads.
Do this experiment $10$ times; what's the largest number of heads that you recorded in those ten runs?


<!-- BEGIN QUESTION -->



```python
def number_of_heads(N):
    """Generate N coin flips and return the number of heads"""
    ...
    return num_heads
```

```python
grader.check("q2")
```

<!-- END QUESTION -->

How likely do you think it is to see a result of at least $600$ heads out of $1000$ flips? Do you think it's somewhat unlikely, very unlikely, or extremely unlikely?

```python
# Put your answer to Problem 2b here

free_response = '...'

print(free_response)
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 5)
```
>>> def test_low_cubes(sum_of_cubes):
...     assert sum_of_cubes(1) == 1
...     assert sum_of_cubes(10) == 3025
...     assert sum_of_cubes(50) == 1625625
>>> test_low_cubes(sum_of_cubes)
```


Tests for q1b (points: 5)
```
>>> def test_low_square_inverse(sum_of_square_inverses):
...     assert sum_of_square_inverses(1) == 1.0
...     assert abs(sum_of_square_inverses(10) - 1.549) < 0.01
...     assert abs(sum_of_square_inverses(20) - 1.596) < 0.01
>>> test_low_square_inverse(sum_of_square_inverses)
```


Tests for q2 (points: 5)
```
>>> def test_num_heads(number_of_heads):
...     assert number_of_heads(10) in list(range(11))
...     assert number_of_heads(100) in list(range(10, 91))
>>> test_num_heads(number_of_heads)
```
